---
title: 'Template Chapter 1'
description: 'This is a template chapter.'
---

## Chapter 1.1 - Coding Exercises

```yaml
type: VideoExercise
key: d666dc7e94
xp: 50
```

`@projector_key`
7ac7ef24515d5c42912be0626e1db076

---

## Insert exercise title here

```yaml
type: NormalExercise
key: 3e6c8e6a47
xp: 100
```

do stuff

`@instructions`
adfadgad

`@hint`
adkfmnadlf

`@pre_exercise_code`
```{r}
rm(list=ls())
```

`@sample_code`
```{r}
x=
```

`@solution`
```{r}
x=1
```

`@sct`
```{r}
ex() %>% check_object("x") %>% check_equal()
success_msg("yay julia!!")
```

---

## Fake News

```yaml
type: VideoExercise
key: c3d4ed3f6f
xp: 50
```

`@projector_key`
f9cccbfdc35a3bcca4bb08c32d13a107
